---
title: 'Week 1'
description: 'Welcome! We will use this platform in the next few weeks for syntax-oriented homeworks, so that you can practice what we have learned in R.'
---

## The author of R

```yaml
type: MultipleChoiceExercise
key: 583049e0db
xp: 50
```

Which R function returns the square root of a number?

`@instructions`
- square
- root
- sqrt
- Sqrt

`@hint`
Remember, R commands are case sensitive!

`@sct`
```{r}
ex() %>% check_mc(correct = 3)
```
